Distinct mechanisms of inhibition of Kv2 potassium channels by tetraethylammonium and RY785

Author response:

The following is the authors’ response to the previous reviews

Public Reviews:

Reviewer #1 (Public review):

Summary:

The authors were seeking to identify a molecular mechanism whereby the small molecule RY785 selectively inhibits Kv2.1 channels. Specifically, the authors sought to explain some of the functional differences that RY785 exhibits in experimental electrophysiology experiments as compared to other Kv inhibitors, namely the charged and non-specific inhibitor tetraethylammonium (TEA). The authors used a recently published cryo-EM Kv2.1 channel structure in the open activated state and performed a series of multi-microsecond-long all-atom molecular dynamics simulations to study Kv2.1 channel conduction under the applied membrane voltage with and without RY785 or TEA present. They observed that while TEA directly blocks K+ permeation by occluding ion permeation pathway, RY785 binds to multiple non-polar residues near the hydrophobic gate of the channel driving it to a semi-closed non-conductive state. They confirmed this mechanism using an additional set of simulations and used it to explain experimental electrophysiology data,

Strengths:

The total length of simulation time is impressive, totaling many tens of microseconds. The authors develop their own forcefield parameters for the RY785 molecule based on extensive QM based parameterization. The computed permeation rate of K+ ions through the channel observed under applied voltage conditions is in reasonable agreement with experimental estimates of the single channel conductance. The authors have performed extensive simulations with the apo channel as well as both TEA and RY785. The simulations with TEA reasonably demonstrate that TEA directly blocks K+ permeation by binding in the center of the Kv2.1 channel cavity, preventing K+ ions from reaching the SCav site. The authors conclude that RY785 likely stabilizes a partially closed conformation of the Kv2.1 channel and thereby inhibits K+ current. This conclusion is plausible given that RY785 makes stable contacts with multiple hydrophobic residues in the S6 helix, which they can also validate using a recently published closed-state Kv2.1 channel cryo-EM structure. This further provides a possible mechanism for the experimental observations that RY785 speeds up the deactivation kinetics of Kv2 channels from a previous experimental electrophysiology study.

Weaknesses:

The authors, however, did not directly observe this semi-closed channel conformation and in fact acknowledge that more direct simulation evidence would require extensive enhanced-sampling simulations beyond the scope of this study. They have not estimated the effect of RY785 binding on the protein-based hydrophobic pore constriction, which may further substantiate their proposed mechanism. And while the authors quantified K+ permeation, they have not made any estimates of the ligand binding affinities or rates, which could have been potentially compared to experiment and used to validate their models.

However, despite those relatively minor weaknesses, the conclusions of the study are convincing, and overall this is a solid study helping us to understand two distinct molecular mechanisms of the voltage-gated potassium channel Kv2.1 inhibition by TEA and RY785, respectively.

Reviewer #2 (Public review):

Summary

In this manuscript, Zhang et al. investigate the conduction and inhibition mechanisms of the Kv2.1 channel, with a particular focus on the distinct effects of TEA and RY785 on Kv2 potassium channels. Using microsecond-scale molecular dynamics simulations, the authors characterize K⁺ ion permeation and RY785-mediated inhibition within the central pore. Their results reveal an inhibition mechanism that differs from those described for other Kv channel inhibitors.

Strengths

The study identifies a distinctive inhibitory mode for RY785, which binds along the channel walls in the open-state structure while still permitting a reduced level of K⁺ conduction. In addition, the authors propose a long-range allosteric coupling between RY785 binding in the central pore and changes in the structural dynamics of Kv2.1. Overall, this is a well-organized and carefully executed study, employing robust simulation and analysis methodologies. The work provides novel mechanistic insights into voltage-gated potassium channel inhibition and may offer useful guidance for future structure-based drug design efforts.

Weaknesses:

The study needs to consider the possibility of multiple binding sites for PY785, particularly given its impact on voltage sensors and gating currents. Specifically, the potential for allosteric binding sites in the voltage-sensing domain (VSD) should be assessed, as some allosteric modulators with thiazole moieties are known to bind VSD domains in multiple voltage-gated sodium channels (Ahuja et al., 2015; Li et al., 2022; McCormack et al., 2013; Mulcahy et al., 2019). Increasing structural and functional evidence supports the existence of multiple ligand-binding modes in voltage-gated ion channels. For example, polyunsaturated fatty acids have been shown to bind to KCNQ1 at both the voltage sensor domain and the pore domain (https://doi.org/10.1085/jgp.202012850). Similarly, cannabidiol has been structurally resolved in Nav1.7 at two distinct sites, one in a fenestration and another near the IFM-binding pocket (https://doi.org/10.1038/s41467-023-39307-6). These advances illustrate that ligand effects cannot always be interpreted based solely on a single binding site identified previously.

Reviewing Editor:

The comments of the reviewers seem thoughtful and constructive. The weaknesses noted in reviews mainly concern mismatch between expectations, created by reading the Abstract, and data in the manuscript. The mismatch could be reconciled by either new simulations examining a semi-open state of the gate and additional RY785 binding sites, or by adjusting wording of the Abstract and Discussion to make it more clear that such simulations were not done.

The Abstract and Discussion have been revised to make clear the computer-simulations presented in our study were designed to specifically validate or refute the hypothesis that RY785 is recognized by the pore domain, not the voltage sensors.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

The authors addressed all the major issues in the original submission identified by the reviewers. I noticed a few minor issues, listed below, which can potentially fix small errors and further improve the readability of the manuscript.

p.3 tetramethyl-ammonium -> tetraethylammonium

p.7 "Snapshot of the final snapshot" -> "Snapshot of the final simulation coordinates"

p. 8 "sigma value" - please spell out what it is.

p. 9 "one or other subunit of the tetramer" -> "one or another subunit of the tetramer" or "one or more subunits of the tetramer"

p 15 "(the net charge of these constructs is thus zero)." -> ""(the net charge of these constructs is zero for these systems)." Please note that using ionizable amino acid residues in their default protonation state does not guarantee net zero charge of the system since the number of cationic and anionic residues is generally not the same.

p. 15 "Two K+ ions were initially positioned in the selectivity filter, one coordinated by residues 373..." Please indicate at which ion binding sites S_1, S_2, e.g. K+ were located and what the residue names are .

SI Figs. S3-S20. Please indicate in the figure captions that all those data are for RY785

SI Fig. S22 and SI Table S1 captions "shown in Fig. S20" -> "shown in Fig. S21"

We thank the Reviewer for this thorough proofreading. We have made the necessary corrections.

Reviewer #2 (Recommendations for the authors):

The authors have addressed most of my comments satisfactorily, with the exception of the first point. Below, I provide further clarification regarding my concern.

First, it appears that the authors may have misunderstood what is meant by the possibility of multiple binding sites for RY785. This does not imply that the central pore is excluded as a binding site. Rather, it refers to the possibility that, in addition to a pore-domain site, the ligand may interact with additional binding sites, either simultaneously or in a statedependent manner. Increasing structural and functional evidence supports the existence of multiple ligand-binding modes in voltage-gated ion channels. For example, polyunsaturated fatty acids have been shown to bind to KCNQ1 at both the voltage sensor domain and the pore domain (https://doi.org/10.1085/jgp.202012850). Similarly, cannabidiol has been structurally resolved in Nav1.7 at two distinct sites, one in a fenestration and another near the IFM-binding pocket (https://doi.org/10.1038/s41467-02339307-6). These advances illustrate that ligand ecects cannot always be interpreted based solely on a single binding site identified previously. Therefore, even if one assumes that there is no precedent for a small-molecule inhibitor that simultaneously acts on both the voltage sensor and pore domain, this does not exclude the possibility that a ligand may bind to both regions in dicerent functional states.

The Reviewer’s opinion came across clearly in the previous version. We however disagree that a computational investigation of the possibility that RY785 binds to the voltagesensors is well-advised at this point, given that the model we propose seemingly ocers a rationale for the inhibitory ecects observed experimentally. Our opinion is also that there is no compelling precedent for the mechanism of inhibition envisaged by the Reviewer – and would argue that neither of the two studies referenced above are compelling examples. As we stated in our previous response to the Reviewer, we believe that the logical next step in this research will be to validate or refute the computational prediction we have put forward, experimentally.

In addition, the present computational study does not provide direct mechanistic evidence to explain the statement that RY785 accelerates voltage-sensor deactivation. Specifically, no simulations were performed to model pore-domain closure or voltage-sensor motion upon RY785 binding. Moreover, alternative binding sites were neither explored nor explicitly excluded, as the simulations only involved placing a single molecule of TEA or RY785 approximately 10 Å below the cytoplasmic gate. Under these conditions, conclusions regarding ecects on voltage-sensor dynamics remain speculative.

That is a fair characterization.

These concerns do not detract from the overall quality of this otherwise strong computational study. There are several straightforward ways to address this issue. For example:

(1) Perform molecular docking or related screening approaches to evaluate potential ligand-binding sites beyond the central pore, particularly in regions proximal to the voltage sensor. This should not impose a substantial additional computational burden for a computational chemistry group.

(2) Revise the abstract and discussion to clarify that the current work focuses exclusively on pore-domain binding and does not explore possible additional binding sites near the voltage sensor. Explicitly stating this limitation would help prevent potential overinterpretation by readers.

We have opted for (2), as noted above.

Changes from the earlier version to the revised version

Abstract

Earlier version: Voltage-gated K+ channels play central roles in human physiology, both in health and disease. A repertoire of inhibitors that are both potent and specific would therefore be of great value, not only as pharmacological agents but also as research tools. The small molecule RY785 has been described as particularly promising in this regard, as it selectively inhibits channels in the Kv2 subfamily with high potency. Kv2 channels are expressed in multiple cell types in humans, and are of particular importance for neuronal function. The mechanism of action of RY785 has not yet been determined at the molecular level, but functional studies indicate it differs from that of less specific inhibitors, such as quaternary-ammonium compounds or aminopyridines; RY785 is distinct also in that it is electroneutral. To examine this mechanism at the single-molecule level, we have carried out a series of all-atom molecular dynamics simulations based on the experimental structure of the Kv2.1 channel in the activated, open state. First, we report a 25-microsecond trajectory calculated in the absence of any inhibitor, under an applied voltage of 100 mV, which demonstrates outward K+ flow under simulation conditions at rates comparable to experimental measurements. Additional simulations in which either RY785 or tetraethylammonium (TEA) is introduced in solution show both inhibitors spontaneously enter the channel through the cytoplasmic gate, with distinct effects. In agreement with prior structural studies, we observe that TEA binds to a site adjacent to the selectivity filter, on the pore axis, thereby blocking the flow of K+ ions. RY785, by contrast, binds to the channel walls, off-axis, and allows K+ flow while the cytoplasmic gate remains open. The observed mode of RY785 binding, however, indicates that its mechanism of action is to stabilize and occlude a semi-open state of the gate, by bridging hydrophobic protein-protein interactions therein; this hypothesis would explain the puzzling experimental observation that RY785 recognition influences the gating currents generated by the voltage sensors, 3 nm away.

Revised version: Voltage-gated K+ channels play central roles in human physiology, both in health and disease. A repertoire of inhibitors that are both potent and specific would therefore be of great value, not only as pharmacological agents but also as research tools. The small molecule RY785 has been described as particularly promising in this regard, as it selectively inhibits channels in the Kv2 subfamily with high potency. Kv2 channels are expressed in multiple cell types in humans, and are of particular importance for neuronal function. The mechanism of action of RY785 has not yet been determined at the molecular level, but functional studies indicate it differs from that of less specific inhibitors, such as quaternary-ammonium compounds or aminopyridines; RY785 is distinct also in that it is electroneutral. To examine this mechanism at the single-molecule level, we have carried out a series of all-atom molecular dynamics simulations based on the experimental structure of the Kv2.1 channel in the activated, open state. First, we report a 25-microsecond trajectory calculated in the absence of any inhibitor, under an applied voltage of 100 mV, which demonstrates outward K+ flow under simulation conditions at rates comparable to experimental measurements. Additional simulations in which either RY785 or tetraethylammonium (TEA) is introduced in solution and allowed to interact with the pore domain show both inhibitors spontaneously enter the channel interior through the cytoplasmic gate, but with distinct effects. In agreement with prior structural studies, we observe that TEA binds to a site adjacent to the selectivity filter, on the pore axis, thereby blocking the flow of K+ ions. RY785, by contrast, binds to the channel walls, off-axis, and allows K+ flow while the cytoplasmic gate remains open. The observed mode of RY785 binding, however, indicates that its mechanism of action is to stabilize and occlude a semi-open state of the gate, by bridging hydrophobic protein-protein interactions therein; this hypothesis would explain the puzzling experimental observation that RY785 recognition influences the gating currents generated by the voltage sensors, 3 nm away.

Introduction

Voltage-gated potassium (Kv) channels are an important class of membrane proteins involved in a variety of crucial physiological processes, both in health and disease. These channels contribute to confer electric excitability to certain types of cells, such as neurons and cardiomyocytes, by sensing and responding to changes in the transmembrane voltage. The nature of this response varies across different types of Kv channels, but in essence it entails a structural mechanism that results in the opening or closing of a pathway across the protein through which K+ ions may flow at a fast rate. This mechanism has been the focus of numerous electrophysiological studies and structure-function analyses spanning several decades, and by now it appears well established (1-4).

Earlier version: As for other proteins of interest, inhibitors of Kv-channel function have been key to advance our understanding of these systems at the molecular level (5). Examples include non-specific blockers, such as quaternary-ammonium compounds (6) and aminopyridines (7); and polypeptide toxins isolated from animal venoms, such as stromatoxin-1 (8), hanatoxin (9) and guangxitoxin (10). A range of small-molecule inhibitors have also been identified (11), but translational applications have been hampered by their modest potency and poor ability to discriminate among different Kv-channel types. In a recent breakthrough, however, a high-throughput screening identified an inhibitor of Kv2 channels, known as RY785, which appears promising both in terms of specificity and potency (12). Subsequent electrophysiological studies of the mode of action of this compound, by Sack and coworkers (13), indicated that RY785 binds to a site in the interior of rat Kv2.1 channels that, at least in part, coincides with that used by tetramethyl-ammonium (TEA), a well-characterized cationic blocker (14-16). Access to this site requires voltage activation of the channel, i.e. opening of the ion-permeation pathway on its cytosolic side. Unlike TEA, though, RY785 is electroneutral, and so it is not immediately apparent whether it impedes K+ flow directly or indirectly. Intriguingly, Sack and coworkers also report RY785 accelerates the deactivation of the voltage sensors (13), which lie over 3 nm away from the TEA binding site. This observation indicates RY785 recognition alters the conformational energetics of the channel, and that its mode of action, as yet unknown, is more complex than that of open-pore blockers such as TEA.

Revised version: As for other proteins of interest, inhibitors of Kv-channel function have been key to advance our understanding of these systems at the molecular level (5). Examples include non-specific blockers, such as quaternary-ammonium compounds (6) and aminopyridines (7); and polypeptide toxins isolated from animal venoms, such as stromatoxin-1 (8), hanatoxin (9) and guangxitoxin (10). A range of small-molecule inhibitors have also been identified (11), but translational applications have been hampered by their modest potency and poor ability to discriminate among different Kv-channel types. In a recent breakthrough, however, a high-throughput screening identified an inhibitor of Kv2 channels, known as RY785, which appears promising both in terms of specificity and potency (12). Subsequent electrophysiological studies of the mode of action of this compound, by Sack and coworkers (13), indicated that RY785 binds to a site in the interior of rat Kv2.1 channels that, at least in part, coincides with that used by tetraethylammonium (TEA), a well-characterized cationic blocker (14-16). Access to this site requires voltage activation of the channel, i.e. opening of the ion-permeation pathway on its cytosolic side. Unlike TEA, though, RY785 is electroneutral, and so it is not immediately apparent whether it impedes K+ flow directly or indirectly. Intriguingly, Sack and coworkers also report RY785 accelerates the deactivation of the voltage sensors (13), which lie over 3 nm away from the TEA binding site. This observation indicates RY785 recognition alters the conformational energetics of the channel, and that its mode of action, as yet unknown, is more complex than that of open-pore blockers such as TEA.

Earlier version: In parallel, Swartz and coworkers recently succeeded to determine an atomic-resolution structure of the channel domain of rat Kv2.1, uninhibited and immersed in a lipid nanodisc, using single-particle cryo-electron microscopy (17). Like other Kv channels, the structure shows Kv2.1 is a homotetramer, and that each subunit includes six trans-membrane alpha-helices (S1-S6). The ion pore is at the center of this complex, formed by assembly of the last two transmembrane helices in each subunit (and the connector in between); the extracellular half of this pore, or selectivity filter, is much narrower than the intracellular half, and features a series of K+ binding sites in a configuration that indicates the filter is in a conductive state. Four voltage-sensing domains surround the pore, each formed by the first four helices in each subunit (17); an additional alpha-helix connects the sensor to the pore in each subunit, thereby coupling the voltage-response of the former to the opening/closing of the latter. Interestingly, the cryo-EM experiment resolved the channel structure in a unique conformation, namely one with the four voltage-sensors in the activated state; accordingly, the intracellular side of the pore domain is wide open. Thus, despite the fact that there is no voltage difference across the lipid nanodisc, the experimental condition somehow mimics the effect of membrane depolarization (i.e. positive-inside voltages), which causes this channel to be predominantly activated/open. To our knowledge, the structure of wild-type Kv2.1 in the deactivated/closed state remains unknown; however, Swartz and coworkers also determined a structure of a Kv2.1 mutant in which pore and voltage domains are seemingly decoupled (17). In this structure the sensors remain activated, and the selectivity filter is unchanged, but the intracellular side of the pore, or gate, is closed; specifically, a rearrangement of helices S5 and S6 causes a hydrophobic constriction to form at the position of residues Ile405 and Pro406, which seems too narrow to permit K+ flux.

Revised version: In parallel, Swartz and coworkers recently succeeded to determine an atomic-resolution structure of the channel domain of rat Kv2.1, uninhibited and immersed in a lipid nanodisc, using single-particle cryo-electron microscopy (17). Like other Kv channels, the structure shows Kv2.1 is a homotetramer, and that each subunit includes six transmembrane alpha-helices (S1-S6). The ion pore is at the center of this complex, formed by assembly of the last two transmembrane helices in each subunit (and the connector in between); the extracellular half of this pore, or selectivity filter, is much narrower than the intracellular half, and features a series of K+ binding sites in a configuration that indicates the filter is in a conductive state. Four voltage-sensing domains surround the pore, each formed by the first four helices in each subunit (17); an additional alpha-helix connects the sensor to the pore in each subunit, thereby coupling the voltage-response of the former to the opening/closing of the latter. Interestingly, the cryo-EM experiment resolved the channel structure in a unique conformation, namely one with the four voltage-sensors in the activated state; accordingly, the intracellular side of the pore domain is wide open. Thus, despite the fact that there is no voltage difference across the lipid nanodisc, the experimental condition somehow mimics the effect of membrane depolarization (i.e. positive-inside voltages), which causes this channel to be predominantly activated/open. To our knowledge, the structure of wild-type Kv2.1 in the deactivated/closed state remains unknown; however, Swartz and coworkers also determined a structure of a Kv2.1 mutant in which pore and voltage domains are seemingly decoupled (17). In this structure the sensors remain activated, and the selectivity filter is unchanged, but the intracellular side of the pore, or gate, is closed; specifically, a rearrangement of helices S5 and S6 causes a hydrophobic constriction to form at the position of residues Ile405 and Pro406, which seems too narrow to permit K+ flux.

In this study, we leverage the discovery of the atomic structure of activated Kv2.1 to gain novel insights into the mode of inhibition of RY785. To do so, we employ atomically-detailed computer simulations of the recognition of this inhibitor by the channel, and quantify its effect on K+ permeation, relative to the uninhibited channel; the results are compared with an analogous analysis of the mode of TEA binding and inhibition, which serve as a control. Taken together, our observations indicate that RY785 inhibits Kv2.1 function not by directly blocking K+ currents, when the pore is open, but by facilitating the closure of its intracellular gate.

Results

Simulation of activated Kv2.1 provides realistic description of outward K+ current

To establish a baseline for this investigation, we first examined the rate and mechanism of K+ permeation through activated Kv2.1 in the absence of any inhibitors. To do so, we carried out an all-atom MD simulation of the channel embedded in a model phospholipid bilayer, under the influence of a transmembrane electric field (Fig. 1). The channel structure employed in this calculation was that recently determined by Swartz and co-workers using cryo-electron microscopy (17); this structure comprises the pore domain and four voltage sensors (Fig. 1), solubilized in a lipid nanodisc. To maximize the statistical significance of our results, the simulation was calculated using an Anton 2 supercomputer, which enabled us to generate a continuous 25-μs trajectory. The simulation was initialized with two K+ ions inside the narrowest portion of the pore, known as selectivity filter, and no voltage difference across the membrane; after 500 ns, a voltage difference of 100 mV was introduced, positive on the inside. As shown in Fig. 2A, at this point K+ ions begin to steadily flow across the channel in a single file and in an outward direction. In the observed mechanism, three K+ ions concurrently reside in the filter, fluctuating in a concerted fashion among five transient binding sites therein, termed S0, S1, S2, S3 and S4 (Fig. 2A, 2B). Arrival of an additional K+ ion through the cytoplasmic gate at a site termed Scav, near the center of a water-filled cavity adjacent to S4, leads to a metastable 4-ion configuration. This configuration is short-lived, on account of the increased electrostatic repulsion between the K+ ions, and the incoming K+ ion often returns to the interior (not shown). An alternative outcome, fostered by the applied electric field, is the ejection of the outermost K+ ion from the filter into the extracellular space. This event is immediately followed by the outward movement of the two other ions within the filter, and that in Scav, thereby restoring the initial 3-ion configuration (Fig. 2B). Four iterations of this molecular cycle thus result in a complete ion permeation event. As has been noted elsewhere, water molecules can co-permeate with K+ but do so rarely; in most of the observed permeation events, the ions become completely dehydrated near the center of the selectivity filter (18, 19).

Structure and simulation of the Kv2.1 channel in the activated state.

(A) Structure of Kv2.1 determined by single-particle cryo-electron microscopy in a lipid nanodisc (17), viewed from the cell intracellular side and along the perpendicular to the membrane. The channel is an assembly of four identical subunits (in colors); the transmembrane ion pore is formed at the center of the assembly, and the voltage sensors are found in the periphery, in a domain-swapped configuration. A K+ ion is shown in magenta, inside the selectivity filter. (B) Simulation system used in this study, comprising the channel (yellow), a phospholipid bilayer (gray), and a 300 mM KCl buffer (magenta, green). The figure depicts the final configuration of the 25-μs MD trajectory described in Figure 2. The total number of atoms is 201,954, some of which are omitted in the figure, for clarity. (C) Close-up of the pore domain and the selectivity filter. Note two protein subunits are omitted, for clarity. The configuration represented is that shown in panel (B); in this configuration, K+ ions are found in sites S1, S3 and S4 within the selectivity filter, while sites S0, S2 and Scav are transiently vacant.

Mechanism of K+ permeation through the Kv2.1 channel.

(A) Time traces of the position of K+ ions along the central axis of the channel as they reach and permeate the selectivity filter towards the extracellular side (black, red and blue). Ions that do not reach the filter are not shown for clarity (those that reach the filter but return to the cytoplasmic side before permeating are shown in orange and cyan). The approximate location of each of the K+ binding sites (S0 through S4, and Scav) along the pore is indicated alongside the plot. (B) Same as (A), for a fragment of the trajectory that illustrates the knock-on mechanism that initiates and completes each of the observed permeation events. (C) Density for K+ ions along the channel axis, relative to the bulk concentration value of 300 mM. Density peaks correspond to each of the K+ binding sites within and adjacent to the selectivity filter. For reference, gray horizontal lines indicate the average position of selected protein atoms: in the selectivity filter these are, from top to bottom, the backbone carbonyl oxygens of residues Y376, G375, V374 and T373, and the sidechain hydroxyl oxygen of T373; and in the cytoplasmic gate, the alpha-carbon of residue P408. All ions in the simulation system contribute to this profile, but only while they reside in a cylindrical volume of diameter equal to 12 Å, centered in and parallel to the channel axis, extending across the whole system.

In total, we observed 45 complete permeation events, i.e. approximately one K+ ion traverses the channel from side to side every 0.5 μs. This high throughput rate appears to be explained by the fact that K+ remains at near bulk density throughout the pore domain (Fig. 2C); as K+ ions cross the channel, they encounter both free-energy wells and free-energy barriers, but it appears that the balance of intermolecular interactions therein is such that these features are relatively shallow, and therefore do not cause large density variations e.g. 100-fold or greater (Fig. 2C). Experimental estimates of the single-channel conductance of Kv2.1 range from 8 to 10 pS depending on the precise condition and type of measurement (20-22); by comparison, the permeation rate observed in our simulation translates into 3 pS. Given the myriad approximations and simplifications inherent to MD simulations, this result is very encouraging, and suggests our methodology results in an adequate description of the mechanism of ion conduction of this channel. We therefore posit, by extension, that the same simulation design should be appropriate to examine how this mechanism might be modulated by an inhibitor, as described below.

Spontaneous binding of tetraethylammonium blocks K+ permeation

Tetraethylammonium (TEA) is a well-characterized non-specific K+-channel blocker; it can inhibit K+ currents completely, whether applied externally or internally, albeit with modest potency. For Kv2.1 in particular, the IC values are around 5 mM and 0.2 mM, respectively (23, 24). Structural and computational studies carried out two decades ago using the bacterial homolog KcsA as a model system revealed that the binding site for internal TEA (and related quaternary-ammonium blockers) is adjacent to Scav (14, 15); from this finding it is logical to infer that internal TEA enters the pore through the cytoplasmic gate and that it precludes K+ ions from reaching the selectivity filter.

To verify this putative inhibitory mechanism, which to our knowledge had not been directly demonstrated before, we carried out a simulation identical to that reported in the previous section, except that TEA was added in solution on the cytoplasmic side of the lipid membrane. The inhibitor was confined to explore a wide cylindrical volume co-axial with the pore but was otherwise permitted to diffuse freely. Under an applied voltage of 100 mV, however, TEA rapidly entered the channel interior through the cytoplasmic gate and did not return to the solution in a 5-μs trajectory (Fig. 3A, 3B). Once in the water-filled cavity under the selectivity filter, the inhibitor continues to be highly dynamic, but tends to reside in a site about 1.5 Å away from Scav (Fig. 3C), and the nitrogen atom that carries the positive charge of the ammonium ion rarely deviates more than 2 Å from the pore axis (Fig. 4A, 4B). Importantly, this simulation did not reveal a single K+ permeation event following TEA binding, in spite of the sustained transmembrane voltage driving K+ outwards (Fig. 3B); the inhibitor prevents cytoplasmic K+ ions from reaching Scav (Fig. 4B), and in doing so, it precludes the onset of knock-on mechanism of permeation through the selectivity filter described in the previous section (Fig. 2B).

Binding of TEA and RY785 to Kv2.1 channel and impact on K+ permeability.

Earlier version: (A) Time trace of the position of TEA (N atom) in the first 500 ns of the simulation, showing its spontaneous binding to the cavity between the selectivity filter and the cytoplasmic gate (marked by the Cα atoms of T373 and P406, respectively, indicated with gray traces). (B) Time trace of the position of TEA for the rest of the 5-µs simulation, alongside those for K+ ions within the selectivity filter, shown as in Figure 2. The location of each of the K+ binding sites therein is indicated alongside the plot. (C) Snapshot of the final snapshot of the simulation, with TEA bound to Scav. Only two of the four channel subunits are shown, for clarity. Residues P406 and I405 are marked in green. (D-F) Same as (A-C), for the simulation of RY785 binding, indicating separately the positions of the central N atom and of the S atom in the distal 5-membered ring.

Revised version: (A) Time trace of the position of TEA (N atom) in the first 500 ns of the simulation, showing its spontaneous binding to the cavity between the selectivity filter and the cytoplasmic gate (marked by the Cα atoms of T373 and P406, respectively, indicated with gray traces). (B) Time trace of the position of TEA for the rest of the 5-μs simulation, alongside those for K+ ions within the selectivity filter, shown as in Figure 2. The location of each of the K+ binding sites therein is indicated alongside the plot. (C) Close-up of the channel in the final snapshot of the simulation, with TEA bound to Scav. Only two of the four channel subunits are shown, for clarity. Residues P406 and I405 are marked in green. (D-F) Same as (A-C), for the simulation of RY785 binding, indicating separately the positions of the central N atom and of the S atom in the distal 5-membered ring.

Bound RY785 does not occlude the K+ pathway, but TEA does.

Earlier version: (A) Probability histograms for the position of TEA and RY785 (in Å) as projected on the plane of the membrane, i.e. perpendicular to the pore, with the origin on its central axis. (B) Snapshot of the simulation of Kv2.1 and TEA described in Fig. 3, with a density map for TEA (blue mesh) calculated for all non-hydrogen atoms and all simulation snapshots. The map is viewed from the membrane plane, on the left, as well as from the cytoplasmic entrance of the pore, on the right. (C) Same as (B), for the simulation of Kv2.1 and RY785. All maps are contoured at the same sigma value (0.04).

Revised version: (A) Probability histograms for the position of TEA and RY785 (in Å) as projected on the plane of the membrane, i.e. perpendicular to the pore, with the origin on its central axis. (B) Snapshot of the simulation of Kv2.1 and TEA described in Fig. 3, with a density map for TEA (blue mesh) calculated for all non-hydrogen atoms and all simulation snapshots. The map is viewed from the membrane plane, on the left, as well as from the cytoplasmic entrance of the pore, on the right. (C) Same as (B), for the simulation of Kv2.1 and RY785. All maps are contoured at the same sigma value (equal to 0.04).

RY785 binding to Kv2.1 permits K+ flow and involves residues in the cytoplasmic gate

As mentioned, RY785 is a specific, potent inhibitor of Kv2.1 (IC50 = 50 nM) (12), but the structural basis for this inhibitory activity is unknown. Unlike TEA, RY785 is electroneutral, and significantly larger, and thus it might have a very different mode of interaction with the channel. To examine this interaction and its mechanistic implications, we carried out an MD simulation identical in every way to that described above for TEA. (Note this simulation required development and optimization of a set of forcefield parameters for RY785, as described in Methods and in Supplementary Information.) This simulation also showed that RY785 can readily penetrate the cytoplasmic gate, though more slowly than TEA (Fig. 3D). While RY785 is not cationic, it features an electric dipole, and the 100-mV voltage difference applied across the simulation system (positive inside) orients the inhibitor as it enters the channel. Specifically, the m-methoxybenzene moiety passes through the gate first and is ultimately closer to the selectivity filter, while the benzimidazole and thiazole moieties, which are more electronegative, cross the gate last and thus end up being closer to the cytoplasmic space (Fig. 3D, 3F). Like TEA, RY785 continued to be very dynamic inside the pore, though it primarily sampled a unique binding mode, in four different configurations reflecting the symmetry of the channel; displacements parallel to the pore axis were far less common than for TEA, and RY785 did not escape from the protein interior in a 5.5-μs trajectory (Fig. 3E, 3F). Nevertheless, we observed that RY785 did not prevent the flow of K+ ions through the channel, nor did it significantly alter the knock-on mechanism within the filter (Fig. 3E). Unlike TEA, RY785 does not prevent cytoplasmic K+ ions from reaching the Scav site, though they clearly do so more rarely; thus the permeation rate we observe is approximately 4 times slower that for the uninhibited channel (Fig. 3E).

Earlier version: Further inspection of the simulation data reveals why RY785 does not block the outward flow of K+ ions. As shown in Fig. 4, this inhibitor binds to the wall of the cavity, through the benzimidazole and thiazole moieties; the m-methoxybenzene ring projects toward the center of the pore, but the high flexibility of this segment leaves an open pathway for K+ throughout the length of the cavity. Logically, this pathway is narrower than for the apo protein, explaining the slower permeation rate, but it is sufficient to permit K+ to access the Scav site and ultimately enter the selectivity filter. This particular mode of interaction stems from persistent hydrophobic interactions with residues in the S6 helices (Fig. 5). In particular, the most important contacts are Val409, Pro406, Ile405, Ile401, and Val398, in one or other subunit of the tetramer; among these, Pro406, Ile401, and Val398 appear to be key, as the inhibitor often contacts these residues simultaneously in two protein subunits (Fig. 5A). As will be discussed below, this observation is important, because several of these residues are very likely to form the hydrophobic seal that closes the cytoplasmic entrance to the pore in the deactivated, closed state of the channel.

Revised version: Further inspection of the simulation data reveals why RY785 does not block the outward flow of K+ ions. As shown in Fig. 4, this inhibitor binds to the wall of the cavity, through the benzimidazole and thiazole moieties; the m-methoxybenzene ring projects toward the center of the pore, but the high flexibility of this segment leaves an open pathway for K+ throughout the length of the cavity. Logically, this pathway is narrower than for the apo protein, explaining the slower permeation rate, but it is sufficient to permit K+ to access the Scav site and ultimately enter the selectivity filter. This particular mode of interaction stems from persistent hydrophobic interactions with residues in the S6 helices (Fig. 5). In particular, the most important contacts are Val409, Pro406, Ile405, Ile401, and Val398, in one or more subunits of the tetramer; among these, Pro406, Ile401, and Val398 appear to be key, as the inhibitor often contacts these residues simultaneously in two protein subunits (Fig. 5A). As will be discussed below, this observation is important, because several of these residues are very likely to form the hydrophobic seal that closes the cytoplasmic entrance to the pore in the deactivated, closed state of the channel.

RY785 binds to the interior wall in the Kv2.1 cavity.

(A) Close-up of the snapshot shown in Fig. 4C, highlighting the most frequently observed protein contacts for RY785; these are on helix S6, which lines the cavity. Only two adjacent protein subunits are shown, for clarity. Note V409, P406 and Ile405 are at the narrowest point of the permeation pathway (aside from the selectivity filter). For reference, the inset shows an alignment of the sequence of this region of S6 (residues 396 to 409) with its equivalent in the KcsA K+ channel (residues 98 to 111), whose closed-state structure is known (25). A more comprehensive comparison of the sequences of the S6 segment in K+ and Ca2+ channels is included in Supp. Fig. 1 (B) Statistical analysis of the interactions depicted in panel (A). A contact was defined as an instance wherein a C atom in the abovementioned hydrophobic sidechains was within 4.5 Å of one of the C/S atoms in RY785; a given snapshot might thus show multiple contacts with each side-chain, sometimes in two different subunits. The plot in the figure quantifies the percentage of simulation snapshots in which a given number of contacts were observed for each side-chain, at minimum; e.g. in 80% of the snapshots RY785 forms ∼5 or more contacts with Pro406; in 1/4 of those, or 20% of the total, the number of observed contacts is ∼10 or more. Contacts with Val409 and Ile405 are mostly in one subunit, while Pro406, Ile401 and Val398 are contacted simultaneously in two subunits.

Bound TEA and RY785 differentially impact K+ access to the selectivity filter

To corroborate the results described above, we used an alternative simulation design wherein no voltage difference was applied across the membrane to drive K+ outwards. Instead, we induced a knock-on event within the selectivity by forcibly driving the central K+ ion (in either S2 or S3) to the S1 site, reasoning this perturbation would result in the immediate ejection of the outermost K+, the movement of the innermost K+ into S2 or S3, and the creation of an artificial vacancy in the S4 site. With either TEA or RY785 occupying the water-filled cavity, just as in the simulations described above, we asked which inhibitor would permit or preclude reloading of site S4 by a K+ ion entering the pore through the cytoplasmic gate.

As shown in Fig. 6, the ion dynamics within the selectivity filter was as anticipated, reproducibly and irrespective of what inhibitor is bound to the channel. The simulations with TEA and RY785 clearly differ otherwise. While TEA is bound, the S4 site remains vacant after the induced knock-on event; only when TEA spontaneously diffuses back into the cytoplasm (note no electric field is applied) prior to the knock-on event, do we observe a K+ ion quickly binding to S4, thereby restoring the preferred 3-ion configuration (Fig. 6A). Reloading of a K+ is also what we observed in all the simulations we carried out with bound RY785, consistent with the observation this inhibitor does not physically occlude the pathway for K+ into the selectivity filter in the open state of the channel (Fig. 6B). These results are completely coherent with those obtained when voltage was used as a driving force for K+ flow in demonstrating that TEA is an open-channel blocker but RY875 is not.

Bound TEA precludes K+ from accessing S4 binding site, but bound RY785 allows it.

(A) Five independent simulations of Kv2.1 bound to TEA wherein a knock-on event was artificially induced in the selectivity filter at t = 100 ns, to create a K+ vacancy in site S4. K+ did not reload this site in the subsequent 100 ns, except when TEA spontaneously dissociated prior to the knock-on event (right). (B) Same as (A), for Kv2.1 bound to RY785. In all simulations, a K+ ion reloads the S4 site within ∼10 ns of the induced knock-on event, corroborating RY785 is not an open-state blocker of Kv2.1.

Discussion

Earlier version: Previous experimental studies have shown that RY785 is a potent, selective inhibitor of Kv2 channels (12). Nonetheless, its site of recognition and precise mechanism of action have been unknown. Allosteric inhibitors featuring thiazole moieties, like RY785 does, have been shown to interact with the voltage-sensing domains of several voltage-gated Na+ channels (26, 27), suggesting RY785 could act similarly. However, available electrophysiological evidence shows that RY785 competes with TEA, a well-studied K+-channel inhibitor known to bind within the pore domain (13); it seems very probable therefore that the binding sites for these two inhibitors coincide, at least in part. The computer simulations of Kv2.1 presented in this study were designed to support or refute this notion without a priori assumptions, but ultimately, the results support that interpretation: when the cytoplasmic gate of the channel is open, both RY785 and TEA can enter the K+ permeation pathway (though not simultaneously), and ultimately dwell in a binding site within the water-filled cavity that precedes the selectivity filter. Aside from this similarity, however, the mechanisms of action of these two inhibitors seem entirely different.

Revised version: Previous experimental studies have shown that RY785 is a potent, selective inhibitor of Kv2 channels (12). Nonetheless, its site of recognition and precise mechanism of action have been unknown. Allosteric inhibitors featuring thiazole moieties, like RY785 does, have been shown to interact with the voltage-sensing domains of several voltage-gated Na+ channels (26, 27), suggesting RY785 could act similarly. However, available electrophysiological evidence shows that RY785 competes with TEA, a well-studied K+-channel inhibitor known to bind within the pore domain (13); it seems very probable therefore that the binding sites for these two inhibitors coincide, at least in part. The computer simulations of Kv2.1 presented in this study were designed to support or refute the hypothesis that RY785 binds to the pore domain, without other a priori assumptions. Ultimately, the results support that notion: when the cytoplasmic gate of the channel is open, both RY785 and TEA can enter the K+ permeation pathway (though not simultaneously), and ultimately dwell in a binding site within the water-filled cavity that precedes the selectivity filter. Aside from this similarity, however, the mechanisms of action of these two inhibitors seem entirely different.

TEA is a small tetrahedral molecule, and is positively charged, though it is larger than a K+ ion. Thus, under positive-inside voltages, it quickly enters the ion permeation pathway and binds near the entrance of the selectivity filter, which is too narrow for TEA to penetrate. Because the molecular symmetry of TEA does not match the 4-fold symmetry of the channel, TEA tends to dwell slightly off-axis; however, this deflection is smaller than the size of a K+ ion, and so on average the blocker appears to be on the channel axis. Consistent with its location and charge, we observe that TEA completely blocks the outward permeation of K+ ions under our simulation conditions; in contrast, in an independent simulation under the same conditions, but in the absence of TEA, we observe instead fast K+ throughput, at a rate comparable to experimental measurements. Comparison of these two simulations shows that TEA blocks K+ currents because it prevents K+ ions in the cytoplasmic space from destabilizing the K+ ions that reside in the selectivity filter.

Earlier version: RY785, however, does not block the flow of K+ ions through the channel, while the cytoplasmic gate is open. RY785 is electroneutral, and mostly hydrophobic; thus, despite being larger than TEA, RY785 does not tend to occupy the center of the pore. Instead, it packs against the walls of the cavity, forming interactions with a cluster of hydrophobic sidechains that, collectively, is specific to Kv2 channels (Fig. 5, Supp. Fig. 1). Recognition of RY785 therefore does not occlude the pathway for K+ into the selectivity filter. How does RY785 inhibit Kv2 channels then? A puzzling but revealing experimental observation is that RY785 recognition by Kv2.1 channels, following activation, accelerates their deactivation (13). That is, RY785 appears to trap itself within the cavity, upon closure of the extracellular gate. Under normal conditions, gate closure requires the repolarization of the membrane, i.e. an energy input. Negative-inside voltages drive the S4 helix in each of the voltage sensors to the down state (creating the so-called gating currents), and this motion is communicated to the pore domain via the S4-S5 linker, ultimately causing a rearrangement in helices S5 and S6, which closes the cytoplasmic gate (Fig. 7). Clearly, from its binding site in the interior of the pore domain, RY785 cannot directly impact the voltage sensors, nor their mechanical coupling with S5-S6. Nonetheless, our simulations indicate an indirect effect. We observe that RY785 forms persistent interactions with the hydrophobic residues that are likely to constrict the pore in the closed state, in the S6 helix (17); in several cases, RY785 interacts simultaneously with the same residue in two adjacent channel subunits. These observations suggest that RY785 might stabilize a semi-closed conformation of the gate that is no longer permeable to K+, even though the structural rearrangement in S5-S6 that is required for full closing has not been achieved completely (Fig. 7). That is, an opening between the four protein subunits would remain at the position of the gate, but this opening would be occupied by RY785, bridging hydrophobic interactions between protein residues therein. Importantly, this hypothetical conformation might not require the voltage sensors to reach the down state completely (Fig. 7); if so, K+ currents would be terminated more readily upon repolarization with bound RY785 than otherwise. This model would also explain the experimental observation that RY785 partially inhibits the currents generated by the motion of the S4 helix, known as gating currents (13), even though S4 is about 3 nm away from the RY785 binding site.

Revised version: RY785, however, does not block the flow of K+ ions through the channel, while the cytoplasmic gate is open. RY785 is electroneutral, and mostly hydrophobic; thus, despite being larger than TEA, RY785 does not tend to occupy the center of the pore. Instead, it packs against the walls of the cavity, forming interactions with a cluster of hydrophobic sidechains that, collectively, is specific to Kv2 channels (Fig. 5, Supp. Fig. 1). Recognition of RY785 therefore does not occlude the pathway for K+ into the selectivity filter. How does RY785 inhibit Kv2 channels then? A puzzling but revealing experimental observation is that RY785 recognition by Kv2.1 channels, following activation, accelerates their deactivation (13). That is, RY785 appears to trap itself within the cavity, upon closure of the extracellular gate. Under normal conditions, gate closure requires the repolarization of the membrane, i.e. an energy input. Negative-inside voltages drive the S4 helix in each of the voltage sensors to the down state (creating the so-called gating currents), and this motion is communicated to the pore domain via the S4-S5 linker, ultimately causing a rearrangement in helices S5 and S6, which closes the cytoplasmic gate (Fig. 7). Clearly, from its binding site in the interior of the pore domain, RY785 cannot directly impact the voltage sensors, nor their mechanical coupling with S5-S6. Nonetheless, our simulations indicate an indirect effect. We observe that RY785 forms persistent interactions with the hydrophobic residues that are likely to constrict the pore in the closed state, in the S6 helix (17); in several cases, RY785 interacts simultaneously with the same residue in two adjacent channel subunits. These observations suggest that RY785 might stabilize a semi-closed conformation of the gate that is no longer permeable to K+, even though the structural rearrangement in S5-S6 that is required for full closing has not been achieved completely (Fig. 7). That is, an opening between the four protein subunits would remain at the position of the gate, but this opening would be occupied by RY785, bridging hydrophobic interactions between protein residues therein. Importantly, this hypothetical conformation might not require the voltage sensors to reach the down state completely (Fig. 7); if so, K+ currents would be terminated more readily upon repolarization with bound RY785 than otherwise. This model would also explain the experimental observation that RY785 partially inhibits the current generated by the motion of the S4 helix, known as gating current (13), even though S4 is about 3 nm away from the RY785 binding site. That is, while this study is not specifically designed to evaluate the possibility that RY785 binds directly to the voltage-sensors, the mechanism we propose would explain the observed effects of RY785 recognition on the gating currents.

Hypothetical mode of Kv2-channel inhibition by RY785.

R785 does not directly block K+ flow; instead, it reshapes the conformational free-energy landscape of the pore domain to stabilize an occluded, partially closed state that does not require S4 to reach the down state in full. The model is an inference based on the simulation data reported in this article and a previous electrophysiological study by Sack and co-workers (13). Note this is a simplified diagram of the functional cycle of the channel, which includes non-conductive or transiently inactivated states, aside from the fully deactivated form (19, 28).

While further studies will be required to fully validate or refute this mechanistic model, recently published work supports our hypothesis. After this study was first reported in 2024, the structure of a Kv2.1 channel with a closed pore domain was successfully determined (29). That structure shows that a Pro-Ile-Pro motif in helix S6 marks the position of the intracellular gate, where the pore domain constricts maximally (aside from the selectivity filter). As illustrated in Fig. 5, this motif is precisely where the benzimidazole and thiazole moieties of RY785 are recognized in our simulations. The inhibition mechanism we outline in Fig. 7 thus seems increasingly plausible. Nonetheless, the fact that both open and closed states of the pore are now known provides an excellent opportunity to examine this model further. Specifically, it is now possible to map the intrinsic conformational free-energy landscape of the pore domain, in a simplified construct lacking the voltage sensors, using enhanced-sampling simulations (30); we anticipate that both open and closed states of the gate will be identifiable free-energy minima in this landscape, and that the closed state will be less favorable than the open state (in absence of the voltage sensors). Comparison of analogous calculations carried out with and without bound RY785, alongside an evaluation of the changes in K+ permeability in each case (31), would provide further quantitative insights into the intriguing mechanism by which RY785 inhibits gating and ion flow in Kv2 channels, and in turn, aid future optimizations of this inhibitor.

Methods

Molecular Dynamics simulations of K+ flux through Kv2.1 in 100 mV

Earlier version: The MD simulations of activated wild-type Kv2.1 are based on the cryo-EM structure of the activated, open state (PDB entry 8SD3) (28). In all cases, the simulations examined a construct encompassing residues 174 to 426 (S1 to S6), with neutral N- and C-termini and all ionizable side -chains in their default protonation state at pH 7 (the net charge of these constructs is thus zero). Two K+ ions were initially positioned in the selectivity filter, one coordinated by residues 373 and 374 and another by residues 375 and 376; a third ion was positioned below the side chain of 373, outside the filter, and two water molecules were modeled between the three ions. (This configuration was very short lived and replaced by others where 3 K+ ions concurrently occupy the filter, typically no water molecules in between – see Results.) For all systems we used Dowser (32) to model structural water molecules within the protein that were not resolved in the experimental density maps. To complete the initial simulation setup of the experimental cryo-EM structure, the construct including ions and water molecules was energy-minimized using CHARMM (33) and the CHARMM36m force field (34-36); specifically, the minimization used the steepest-descent algorithm, for 250 steps, and then conjugate-gradient algorithm, for another 250 steps.

Revised version: The MD simulations of activated wild-type Kv2.1 are based on the cryo-EM structure of the activated, open state (PDB entry 8SD3) (28). In all cases, the simulations examined a construct encompassing residues 174 to 426 (S1 to S6), with neutral N- and C-termini and all ionizable side-chains in their default protonation state at pH 7, which in this construct results in a net charge of zero. Two K+ ions were initially positioned in the selectivity filter, one coordinated by residues 373 and 374 (site S3) and another by residues 375 and 376 (site S1) (Fig. 1C); a third ion was positioned below the side chain of 373, outside the filter, and two water molecules were modeled between the three ions. (This configuration was very short lived and replaced by others where 3 K+ ions concurrently occupy the filter, typically no water molecules in between – see Results.) For all systems we used Dowser (32) to model structural water molecules within the protein that were not resolved in the experimental density maps. To complete the initial simulation setup of the experimental cryo-EM structure, the construct including ions and water molecules was energy-minimized using CHARMM (33) and the CHARMM36m force field (34-36); specifically, the minimization used the steepest-descent algorithm, for 250 steps, and then conjugate-gradient algorithm, for another 250 steps.

The channel was simulated in a POPC lipid bilayer and a 300 mM KCl solution. (Note that while this concentration is atypical in the intracellular environment, it is often used in electrophysiological studies. In the simulation context, this choice is sensible in that it leads to a greater number of permeation events, and therefore, greater statistical confidence in the calculated ion conductance, or its inhibition.) To generate a molecular model of this membrane/solvent environment, we created a coarse-grained (CG) POPC-lipid bilayer in 300 mM KCl in an orthorhombic box of ∼150 × 150 × 100 Å using insane.py (37). To equilibrate this CG system, we carried out a 20-μs MD simulation using GROMACS 2018.8 (38) and the MARTINI 2.2 force field (39, 40) at constant temperature (303 K) and constant semi-isotropic pressure (1 atm) and with periodic boundary conditions. The integration time-step was 20 fs. We then embedded the structure of the Kv2.1 channel in this environment; to do so, the atomic structure of the channel was first coarse-grained with martinize.py (39), and overlaid onto the equilibrated membrane/solvent system, removing all overlapping lipid and water molecules. Then, to optimize the resulting protein/lipid/solvent interfaces, we carried out a 10-μs MD simulation of the complete system using GROMACS 2018.8 (38) and MARTINI 2.2 (39, 40) at constant temperature (303 K) and constant semi-isotropic pressure (1 atm) and with periodic boundary conditions and an integration time step of 20 fs. Having ascertained the equilibration of the membrane structure and of the position and orientation of the protein in the lipid bilayer, the final snapshot of the CG trajectory was transformed into an all-atom representation compatible with the all-atom CHARMM36m force field (34-36). To do so, lipid and solvent molecules were back-mapped onto all-atom models (41), while the CG version of Kv2.1 was replaced (not back-mapped) with the energy-minimized all-atom construct described above, after optimally superposing the Cα trace of the latter onto that of the former. The resulting all-atom molecular system includes 526 POPC lipids, 38,629 water molecules, and 208 K+ and 239 Cl− (300 mM KCl) in an orthorhombic box of ca. 150 × 150 × 100 Å. To further optimize this all-atom model, the simulation system was first energy-minimized for 5,000 steps with NAMD 2.13 (42, 43) and CHARMM36m (34-36), using the conjugate-gradient algorithm. We then carried out a series of MD simulations wherein structural restraints applied to the protein and ions/water within the selectivity filter are progressively weakened over the course of ∼150 ns, as previously described (19, 28). These simulations were carried out using NAMD 2.13 (42, 43) and CHARMM36m (34-36) at constant temperature (298 K) and constant semi-isotropic pressure (1 atm) with periodic boundary conditions and an integration time-step of 2 fs. Electrostatic interactions were calculated using with the Particle-Mesh Ewald method (44), with a real-space cutoff value of 12 Å; van der Waals interactions were also cut off at 12 Å, with a smooth switching function taking effect at 10 Å.

To quantify the ion-conducting properties of activated Kv2.1, a 25-μs MD trajectory was calculated under an applied voltage (positive inside) using an ANTON2 supercomputer (45) and the all-atom CHARMM36m force field (34-36). The starting configuration for this simulation was the final configuration of the abovementioned equilibration process. The simulation was again carried out at constant temperature (298 K) and semi-isotropic pressure (1 atm), in this case set with the Nose-Hoover thermostat (46, 47) and the Martyna-Tobias-Klein barostat (48), respectively, and with periodic boundary conditions and an integration time-step of 2.5 fs. Electrostatic interactions were calculated using the Gaussian-split Ewald method (49); van der Waals interactions were cut off at 10 Å. To preclude large-scale changes in the secondary or tertiary structure of the channel that might develop in the 25-μs timescale, due to cumulative forcefield inaccuracies, the energy function of the simulation was supplemented with a weak biasing potential that favors the known experimental geometry, while permitting the structure to locally fluctuate as required for ligand recognition and ion permeation (19, 28). This potential acts on all ϕ, ψ and χ1 dihedral angles in the protein, and is defined by: where θt is the value of each dihedral angle at time t in the simulation, θexpt denotes the corresponding value in the experimental structure, and k = 1 kBT. Note this potential is identical to that used in previous studies based on multi-microsecond ANTON2 simulations (50-52), except the magnitude of the bias in this study is considerably weaker (i.e. smaller k). To drive K+ ions outward, a transmembrane voltage difference was applied across the membrane, positive inside. This voltage difference was applied as a jump from 0 to 100 mV at timepoint 0.5 μs and sustained thereafter. The desired voltage resulted from application of an outwardly directed, constant electric field perpendicularly to the membrane plane (53). For a simulation box of approximately 92 Å in that direction, a voltage difference of 100 mV corresponds to an electric field of 0.024 kcal mol−1 Å−1 e−1 (note 1 kcal mol−1 Å−1e−1 = 43.4 mV/Å).

Simulations TEA and RY785 binding and inhibition in 100 mV

Additional MD simulations were carried out to examine the mode of TEA and RY785 binding and inhibition. These simulations used as starting point the final configuration of the 25-μs MD trajectory of described above; a molecule of TEA or RY785 was added about 10 Å below the cytoplasmic gate, and overlapping water molecules were removed. These constructs were then energy-minimized using CHARMM (33) and the CHARMM36m force field (34-36); specifically, the minimization consisted of 100 steps using the steepest-descent algorithm with the protein and K+ ions within fixed. To further optimize these all-atom models, the simulation systems were then energy-minimized for 5,000 steps with NAMD 2.13 (42, 43) and the CHARMM36m force field (34-36), using the conjugate-gradient algorithm. Again using NAMD 2.13, we then carried out a 20-ns MD simulation of each channel-inhibitor system wherein weak structural restraints identical to those used in the 25-μs MD trajectory were applied to the protein. In addition, a flat-bottom confining potential was applied to TEA or RY785 to keep the inhibitor inside a cylindrical volume co-axial with the channel pore, 20 Å in diameter. Specifically, we used atoms Y376:Cβ and V409:Cβ in the four protein subunits to define this axis, and the inhibitor was confined by a potential of the form: where dt is the distance between the center-of-mass of the inhibitor and the channel axis, perpendicularly to that axis, do is 10 Å and k = 100 kcal/mol/Å2. To preclude the inhibitor from diffusing across the periodic boundary of the system (i.e. preserving it in the cytoplasmic side of the channel), a confining potential of the same functional form was also applied relative to the center of the four V409:Cβ atoms; in this case dt is the projection of the distance onto the channel axis, while do is 15 Å. Following these 20-ns equilibrations, trajectories of 5-μs and 5.5-μs were calculated on ANTON2 for the TEA and RY785 systems respectively, using exactly the same settings as those reported above for the channel-only simulation.

Simulations TEA and RY785 inhibition after induced K+ knock-on event

The simulations described in Fig. 6 used as starting point the final configurations of the trajectories calculated with ANTON2 for the channel-inhibitor systems (Fig. 3). In these starting configurations, three K+ ions occupy the selectivity filter, and TEA or RY785 resides within the cavity. To create a vacancy in the S4 site, without a voltage difference across the membrane, the central K+ ion (which fluctuates between S2 and S3) was driven outwards at t = 100 ns, by activating a set of restraints that bring the ion within coordination distance of the carbonyl O atoms of Y376 and G375 in all four subunits (ca. 3 Å, k = 3 kcal/mol/Å2). As a result, the outermost K+ ion (which fluctuates between S0 and S1) was ejected, and the innermost K+ ion moved to the center of the filter, creating the desired vacancy in S4. The simulations were then continued for 100 ns to examine whether or not K+ ions would reload this vacancy. These simulations were carried out using NAMD 3.0b6 (42, 43) and the CHARMM36m force field (34-36). All other settings, including confining potentials, were as described above.

Development of molecular-mechanics forcefield for RY785

To our knowledge, no molecular-mechanics (MM) forcefield for RY785 was publicly available prior to this work. Thus, a new forcefield had to be developed and optimized to be able to introduce this molecule in our simulations of Kv2.1 – specifically one compatible with the CHARMM General Force Field, or CGenFF (54). Following recommended procedures (54-56), we carried out a series of quantum chemical calculations to generate ab-initio benchmarking data, including atomic charges, optimized molecular geometries and multiple potential energy surfaces (PES) generated by relaxed-scan of selected bond- and dihedral-angles; these calculations were performed in the gas phase using Gaussian 16 (Gaussian, Inc., Wallingford CT, 2016), using the MP2/6-31G(d) method. Structural properties and calculated PESs were then used to adjust atomic-charge, bond-angle and dihedral-angle parameters lacking adequate entries in CGenFF. These MM parameters were refined targeting the benchmark QM data, including interaction energies and H-bond distances for multiple RY785-H2O binary complexes. Each of these RY785-H2O conformers probes the interaction of a H2O molecule with a different H-bond donor (C-H, or N-H) or acceptor (N or S) in RY785. The complexes were set up with an ideal H-bond interaction between RY785, in its MP2/6-31G(d) optimized geometry, and a water molecule, in its idealized TIP3P geometry (57). The interaction distances in these complexes were optimized at the HF/6-31G(d) level while keeping all other degrees of freedom fixed; the interaction energy was then calculated without correction for basis set superposition error and multiplied by a factor 1.16, as suggested elsewhere (56). Detailed results are provided in Supplementary Information.